[1613557711] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.001 second response time
[1613557771] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.001 second response time
[1613557824] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;1;connect to address 127.0.0.1 and port 22: Connection refused
[1613557831] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;3;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.001 second response time
[1613557862] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613557884] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;2;connect to address 127.0.0.1 and port 22: Connection refused
[1613557891] SERVICE ALERT: localhost;HTTP;WARNING;HARD;4;HTTP WARNING: HTTP/1.1 401 Unauthorized - 691 bytes in 0.000 second response time
[1613557922] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613557945] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;3;connect to address 127.0.0.1 and port 22: Connection refused
[1613557982] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;3;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613558005] SERVICE ALERT: localhost;SSH;CRITICAL;HARD;4;connect to address 127.0.0.1 and port 22: Connection refused
[1613558042] SERVICE ALERT: localhost;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1613571910] SERVICE ALERT: ubuntu;HTTP;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 403 Forbidden - 1530 bytes in 0.006 second response time
[1613571970] SERVICE ALERT: ubuntu;HTTP;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 403 Forbidden - 1530 bytes in 0.002 second response time
[1613572030] SERVICE ALERT: ubuntu;HTTP;WARNING;SOFT;3;HTTP WARNING: HTTP/1.1 403 Forbidden - 1529 bytes in 0.002 second response time
[1613572093] SERVICE ALERT: ubuntu;HTTP;OK;SOFT;4;HTTP OK: HTTP/1.1 200 OK - 50509 bytes in 2.400 second response time
[1613574179] SERVICE ALERT: viewnext104;HTTP_JENKINS;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 401 Unauthorized - 791 bytes in 0.002 second response time
[1613574299] SERVICE ALERT: viewnext104;HTTP_JENKINS;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 401 Unauthorized - 791 bytes in 0.003 second response time
[1613574373] SERVICE ALERT: viewnext104;HTTP_JENKINS;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 401 Unauthorized - 791 bytes in 0.003 second response time
[1613574703] SERVICE ALERT: viewnext104;HTTP_JENKINS;OK;HARD;3;HTTP OK: HTTP/1.1 200 OK - 48079 bytes in 0.440 second response time